## Lab 0 - Intro Tutorial

This is the lab 0 notebook! It contains instructions for an icebreaker activity, introduces the format of subsequent lab assignments, and serves as practice for coding using Python in Jupyter Notebooks.

## Icebreaker activity:
(~7 minutes) 
* Get into your groups and figure out five things you have in common. You will share the most interesting thing you have in common when we reconvene. 
* Create a team name using an adjective + an animal (for example, Voracious Viper.)  If you need some suggestions, you can use this [random adjective generator](https://randomwordgenerator.com/adjective.php) and/or [animal generator](https://www.generatorslist.com/random/animals/random-animal-generator).

## Notebook Intro

Here, you'll get some practice answering questions and validating your answers using error handling.

```python
import numpy as np
from numpy.testing import assert_raises, assert_equal
```

In the cell below, write a function called ``num_squares`` that takes in one integer parameter $n$ and returns the square $n^2$.


<!-- grade_id: num_squares_function, points: 0, grade: false, solution: true -->
```python
def num_squares(n):
    ### BEGIN SOLUTION
    return n**2
    ### END SOLUTION
```

<!-- grade_id: num_squares_test1, points: 1, grade: true, solution: false -->
```python
# (1pt.) Run cell to check that num_squares works for a few values. 
assert_equal(num_squares(2),4)
assert_equal(num_squares(-5),25)
assert_equal(num_squares(0),0)
```

In the cell below, write a function ``arr_func(n)``. The function should print an array with entries spanning the integers between $1$ and $n$ and return that array.  The arrray should be named ``arr``.

<!-- grade_id: cell-892448a2483ca4bc, points: 0, grade: false, solution: true -->
```python
# (1 pt) 
def arr_func(n):
    ### BEGIN SOLUTION
    arr = np.array(range(1,n+1))
    return arr
    ### END SOLUTION
```

<!-- grade_id: arr_functest, points: 1, grade: true, solution: false -->
```python
#  Run this cell to check that arr_func returns the correct array
assert_equal(arr_func(5), [1,2,3,4,5])
assert_equal(arr_func(10), [1,2,3,4,5,6,7,8,9,10])
```

In the cell below, write a function ``val_err(n)`` that raises a ``ValueError`` if the input $n$ is $<1$, and returns $n$ otherwise.
Hint: add the following code.

``
if n < 1:
    raise ValueError("n must be greater than or equal to 1 ")
``

Test this code out, and make sure you get the right error!

<!-- grade_id: value_error, points: 0, grade: false, solution: true -->
```python
# (1 pt)
def val_err(n):
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1 ")
    return n
    ### END SOLUTION
```

<!-- grade_id: valueerror_test, points: 1, grade: true, solution: false -->
```python
# Run this cell to check if the function raises desired ValueError
assert_raises(ValueError, val_err, -1)
assert_equal(val_err(5),5)
```

In the cell below, write a function ``squares`` that takes in an integer $n$. The function should create an array of integers spanning the interval $1$ to $n$ and return a *list* containing the squares of each entry of that array. For example, if $n=5$, the array should return the list ``[1 4 9 16 25]``. There are many ways to accomplish this task; for example, using list comprehension, loops, or NumPy methods.  You may use any method as long as it returns the correct list! Don't worry about the case $n=0$.

<!-- grade_id: squares_def, points: 0, grade: false, solution: true -->
```python
# (1 pt)
def squares(n):
    ### BEGIN SOLUTION
    da_list = []
    for i in range(n):
        da_list.append((i+1)*(i+1))
    return da_list
    ### END SOLUTION
```

In the cell below, create a variable called ``a`` and set it equal to ``squares(5)``.

```python
a = squares(5)
```

<!-- grade_id: squares_test, points: 1, grade: true, solution: false -->
```python
# Run this cell to check that your squares(n) function returns the correct list. 
assert_equal(a, [1,4,9,16,25])
```

To submit lab 0: save this notebook as a PDF using the print feature of your browser and submit to "Lab 0: Intro to Notebooks" on the Assignments page on BruinLearn.

```python
"For Question 1, I asked ChatGPT: Write a function called num_squares that takes in one integer parameter n and returns the square n^2"

"Here is ChatGPT's response:"

def num_squares(n):
    return n ** 2

# This function uses the exponentiation operator (**) to calculate the square of the input integer n. Feel free to try it out with different values!"

# Example usage:
result = num_squares(5)
print(result)

# I give this a 4/4 rating for usefulness, accuracy, relevance, and details
```

Output:
```
25
```

```python
"""For Question 2, I asked ChatGPT: Write a function arr_func(n). The function should print an array with entries spanning the 
integers between 1 and n and return that array. The array should be named arr."""

"Here is ChatGPT's response:"

def arr_func(n):
    arr = list(range(1, n+1))
    print(arr)
    return arr

# This function uses the range function to generate a list of integers from 1 to n (inclusive). It then prints the array and returns it. 
# Feel free to try it out with different values for n!

# Example usage:
result_array = arr_func(5)

# I give this a 4/4 rating for usefulness, accuracy, relevance, and details.
```

Output:
```
[1, 2, 3, 4, 5]
```

```python
"""For Question 3, I asked ChatGPT: Write a function val_err(n) that raises a 
ValueError if the input n is less than 1, and returns n otherwise."""


"Here is ChatGPT's response:"

def val_err(n):
    if n < 1:
        raise ValueError("Input must be greater than or equal to 1")
    return n

# In this function, an if statement checks if n is less than 1. If it is, a ValueError is raised with a corresponding error message. 
# Otherwise, the function returns n. Feel free to test it with different values for n!

# Example usage:
try:
    result = val_err(5)
    print(result)
except ValueError as e:
    print(e)

# I give this a 4/4 rating for usefulness, accuracy, relevance, and details
```

Output:
```
5
```

```python
"""For Question 4, I asked ChatGPT: Write a function squares that takes in an integer n. The function should create an array of integers spanning the 
interval 1 to n and return a list containing the squares of each entry of that array. Don't worry about the case n = 0."""

"Here is ChatGPT's response:"

def squares(n):
    arr = list(range(1, n+1))
    square_list = [x ** 2 for x in arr]
    return square_list

# This function uses a list comprehension to calculate the square of each entry in the array generated by range(1, n+1). 
# The resulting list of squares is then returned. Feel free to try it out with different values for n!

# Example usage:
result_squares = squares(5)
print(result_squares)

# I give this a 4/4 rating for usefulness, accuracy, relevance, and details
```

Output:
```
[1, 4, 9, 16, 25]
```